```python
# Initialize Otter
import otter
grader = otter.Notebook("assignment03.ipynb")
```

## Assignment 3

### Do not change the code that has been provided to you. You can add code but do not change/remove any components provided by us.

```python
import numpy as np
```

**Question 1:** Complete the Class `Stack` and Class `Queue` that provides functionality based on rules given below.<br>

### Stack
<ul>
    <li>Initialize stack with an empty list.</li>
    <li>Implement class functions based on the comments given in them.</li>
</ul>

### Queue
<ul>
    <li>Initialize queue with an empty list.</li>
    <li>Implement class functions based on the comments given in them.</li>
</ul>



```python
class Stack:
    def __init__(self):
        '''
        Initialize stack with empty list
        '''
        self.stack = None
        ...
    def display(self):
        print(self.stack)
    def size(self):
        '''
        Returns the size of the stack
        '''
        ...
    def push(self,value):
        '''
        Pushes value to the top of the stack
        Ex: stack = [1,2,3,4]
        push(5) -> stack = [5,1,2,3,4]
        '''
        ...
    def pop(self):
        '''
        Pops value from the top of the stack
        Ex: stack = [1,2,3,4]
        pop() -> stack = [2,3,4]
        
        return the popped value
        '''
        top = None
        ...
        return top
```

```python
class Queue:
    def __init__(self):
        '''
        Initialize queue with empty list
        '''
        self.queue = None
        ...
    def display(self):
        print(self.queue)
    def size(self):
        '''
        Returns the size of the queue
        '''
        ...
    def push(self,value):
        '''
        Pushes value to the back of the queue
        Ex: queue = [1,2,3,4]
        push(5) -> queue = [1,2,3,4,5]
        '''
        ...
        return self.queue
    def pop(self):
        '''
        Pops value from the front of the queue
        Ex: queue = [1,2,3,4]
        pop() -> queue = [2,3,4]
        
        return the popped value
        '''
        front = None
        ...
        return front
```

```python
grader.check("q1")
```

**Question 2:** Complete the Class `LinkedList` and Class `Tree` that provides functionality based on rules given below.<br>

### LinkedList
<ul>
    <li>Implement class functions based on the comments given in them.</li>
</ul>

### Tree
<ul>
    <li>Implement class functions based on the comments given in them.</li>
</ul>



```python
'''
These nodes will be instrumental in implementing LinkedList and Tree data structures.
'''
class LLNode:
    def __init__(self,value):
        self.val = value
        self.next = None
class TreeNode:
    def __init__(self,value):
        self.val = value
        self.left = None
        self.right = None
```

```python
class LL:
    def __init__(self,val):
        '''
        Initialize the Linked List with a node with value : val
        '''
        self.head = None
        ...
    def print(self):
        '''
        Print the linked list from head to the last node
        '''
        ...
    def get_list(self):
        '''
        Return List form of the Linked List
        LL = 1->2->4->3->None
        return [1,2,4,3]
        '''
        l=[]
         ...
        return l
    def add(self,val):
        '''
        Add a LLNode to the end of the existing Linked List
        Ex :
        If the Linked list is at a state : 1->2->4->3->None
        you have self.head (1) and a value (5)
        You need to return the head of the LinkedList which is : 1->2->4->3->5->None
        '''
        ...
        return self.head
    def delete(self, value):
        '''
        Delete a LLNode with val = value in the Linked List
        Ex :
        If the Linked list is at a state : 1->2->4->3->None
        you have self.head (1) and value (4)
        You need to return the head of the LinkedList which is : 1->2->3->None
        '''
        ...
        
```

```python
'''
Note: In the following question, you will need queue to complete a few functions, use the queue we created previously in this assignment.
'''
class Tree:
    def __init__(self):
        self.root = None
        self.inorder = []
        self.preorder = []
        self.postorder = []
        self.levelorder = []
    def insert_val(self,value):
        '''
        Insert a new TreeNode with val = value in the tree.
        The idea is to traverse through the tree to find a node whose left child is empty, and make the new Tree Node the left child of this node.
        If all the left childs are full, we find a node whose right child is empty, we make the new Tree node as the right child of this node.
        '''
        ...
        
    def height(self, node):
        ...
    def level_order(self):
        '''
                       1
                    2     3
                  4  5   6  7
        levelOrder : [1,2,3,4,5,6,7]
        
        '''
        ...
    def inorder_traversal(self,root):
        '''
                       1
                    2     3
                  4  5   6  7
        Inorder : [4,2,5,1,6,3,7]
        algorithm:
        inorder(left)
        visit node
        inorder(right)
        '''
        ...
    def preorder_traversal(self,root):
        '''
                       1
                    2     3
                  4  5   6  7
        preorder : [1,2,4,5,3,6,7]
        algorithm:
        visit node
        preorder(left)
        preorder(right)
        '''
        ...
    def postorder_traversal(self,root):
        '''
                       1
                    2     3
                  4  5   6  7
        postorder : [4, 5, 2, 6, 7, 3, 1]
        algorithm:
        postorder(left)
        postorder(right)
        visit node
        '''
        ...
    
```

```python
grader.check("q2")
```

**Question 3:** Design a fast food chain ordering system.<br>

### Based on the the instructions given below design a Fast Food chain ordering system.

<ul>
    <li>You are provided 4 classes: Food class, Drink class, Order class and OrderProcessor class. You need to fill in the methods to make the ordering system function as required. </li>
    <li>Each method has a set of comments which will tell you what is expected as an output of that method.
    <li>You will need to maintain a queue in the OrderProcessor task to optimize order times based on time taken to prepare the order. Use the Queue datastructure you constructed above for this operation. You are also provided with reference dictionaries to calculate estimated preparation time and estimated cost.</li>
    <li>Given the complexity of this task, we will not have any hidden test case for this question.</li>
</ul>


Hint: We highly recommend using one of the data structure and its methods we contructed in the previous questions

```python
'''
We have provided OrderClass.py along with this notebook. You need to complete Order, Food and Drink class in OrderClass.py file.
We will then import those classes here and use them in OrderProcessor.
'''

from OrderClass import Food,Drink,Order
```

```python
class OrderProcessor:
    def __init__(self):
        '''
        Initialize a queue to handle incoming orders. 
        Note: We will use self.time to track the amount of time elapsed.
        '''
        self.time = 0
        self.queue = None
        ...
    def timeElapsed(self, time):
        self.time = time+self.time

    def addOrder(self,items):
        '''
        Given a list of items, create an order object.
        Example of list of items:
        [
            {
                'item_type':'Food',
                'item_name':'Fried',
                'item_size':2
            },
            {
                'item_type':'Drink',
                'item_name':'Coke',
                'item_size':'small'
            }
        ]
        Note: The order should have an order id based on the number of orders taken that day. That means that if the queue is empty, the order_id of the first order is 'order_1'.
        If the queue is at this state :[Order(order_id = 'order_5'), Order(order_id = 'order_7'), Order(order_id = 'order_11')]. The order_id of the next order should be 'order_12',
        assuming that the last order that was added had order_id = 'order_11'
        Adds an order to the Queue based on the prep time of the order and the time it arrived (timestamp here refers to arrival_time)
        Example. 
        Queue = [Order1(prep_time = 5, timestamp = 2), Order2(prep_time = 12, timestamp = 1), Order3(prep_time = 15, timestamp = 4)]
        Now if we want to add this order to the queue: Order4(prep_time = 7, timestamp = 5), then the queue after adding the order should look like this:
        Queue = [Order1(prep_time = 5, timestamp = 2), Order4(prep_time = 7, timestamp = 5), Order2(prep_time = 12, timestamp = 1), Order3(prep_time = 15, timestamp = 4)]
        '''
        ...
    def displayQueue(self):
        '''
        return a list of order_ids of the orders in the queue as the status of the queue
        Example.
        Queue = [Order1(prep_time = 5, timestamp = 2), Order4(prep_time = 7, timestamp = 5), Order2(prep_time = 12, timestamp = 1), Order3(prep_time = 15, timestamp = 4)]
        return [order_1, order_4, order_2, order_3]
        '''
        status = []
        ...
        return status
    def deliverOrder(self):
        '''
        Deliver an order if it is ready, i.e. self.time > prep time + timestamp (arrival_time) for that order.
        You need to return a list of lists that specify the following:
        Example:
        self.time = 12
        Queue = [Order1(prep_time = 5, timestamp = 2), Order4(prep_time = 7, timestamp = 5), Order2(prep_time = 12, timestamp = 1), Order3(prep_time = 15, timestamp = 4)]
        then you need to deliver Order1 and Order 4
        return [
        [order_id of Order1, total cost of Order1],
        [order_id of Order4, total cost of Order4],
        ]
        '''
        output_list = []
        ...
        return output_list
```

```python
'''
Here is a small example so you understand the implementation
'''
orderProcessor = OrderProcessor()
print(f'Queue: {orderProcessor.displayQueue()}')
orderProcessor.timeElapsed(1)
order_1 = [{'item_type':'Food','item_name':'Fries','item_size':2}, ## cost 9.99 * 2 = 19.98 
           {'item_type':'Drink','item_name':'Coke','item_size':'small'}] ## cost 3.99
orderProcessor.addOrder(order_1) # total prep time = 5(fries) + 5(fries) + 1 (coke) = 11
print(f'Queue: {orderProcessor.displayQueue()}')
orderProcessor.timeElapsed(3)
order_2 = [{'item_type':'Food','item_name':'Slider','item_size':1}] ## cost 10.99
orderProcessor.addOrder(order_2) # total prep time = 6
print(f'Queue: {orderProcessor.displayQueue()}')
orderProcessor.timeElapsed(2)
order_3 = [{'item_type':'Drink','item_name':'Pepsi','item_size':'small'}] ## cost 3.99
orderProcessor.addOrder(order_3) # total prep time = 1 (coke) = 1
print(f'Queue: {orderProcessor.displayQueue()}')
print(f'Time Elapsed: {orderProcessor.time}')
orderProcessor.timeElapsed(2)
print(f'Time Elapsed: {orderProcessor.time}')
print(f'Delivery: {orderProcessor.deliverOrder()}')
print(f'Queue: {orderProcessor.displayQueue()}')
orderProcessor.timeElapsed(2)
print(f'Time Elapsed: {orderProcessor.time}')
print(f'Delivery: {orderProcessor.deliverOrder()}')
print(f'Queue: {orderProcessor.displayQueue()}')
orderProcessor.timeElapsed(2)
print(f'Time Elapsed: {orderProcessor.time}')
print(f'Delivery: {orderProcessor.deliverOrder()}')
print(f'Queue: {orderProcessor.displayQueue()}')

## Output:
'''
Queue: []
Queue: ['order_1']
Queue: ['order_2', 'order_1']
Queue: ['order_3', 'order_2', 'order_1']
Time Elapsed: 6
Time Elapsed: 8
Delivery: [['order_3', 3.99]]
Queue: ['order_2', 'order_1']
Time Elapsed: 10
Delivery: [['order_2', 10.99]]
Queue: ['order_1']
Time Elapsed: 12
Delivery: [['order_1', 23.97]]
Queue: []
'''
```

```python
grader.check("q3")
```

### Submission Guidelines: 
#### 1. Save and run all the cells of this notebook.
#### 2. All the assignments in this course will have two types of test cases, local test cases and hidden test cases. You will be able to test your work on local test cases by running the cell  with - grader.check_all() or grader.check("question number")
#### 3. Once you are happy with the output of grader.check_all(), you should save the following files in a zip to gradescope:  assignment03.ipynb, OrderClass.py  (Make sure the name of the files is consistent when you upload it to gradescope)
#### 4. You will be able to see the result of your work on hidden tests once the grades are released for a given assignment.

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Tests for q1 (points: 20)
```
>>> stack = Stack()
>>> stack.push(2)
>>> stack.push(3)
>>> stack.push(4)
>>> top=stack.pop()
>>> top=stack.pop()
>>> stack.push(2)
>>> stack.stack == [2,2]
True
```
```
>>> def test():
...     queue = Queue()
...     queue.push(2)
...     queue.push(3)
...     queue.push(4)
...     front=queue.pop()
...     front=queue.pop()
...     queue.push(2)
...     return queue
>>> queue = test()
>>> queue.queue == [4,2]
True
```


Tests for q2 (points: 20)
```
>>> def LLTest():
...     ll = LL(1)
...     ll.add(2)
...     ll.add(3)
...     ll.add(4)
...     ll.delete(4)
...     ll1 = LL(1)
...     ll1.delete(1)
...     return ll,ll1
>>> ll,ll1 = LLTest()
>>> ll.get_list() == [1,2,3] and ll1.get_list() == []
True
```
```
>>> def TreeTest():
...     tree = Tree()
...     tree.insert_val(1)
...     tree.insert_val(2)
...     tree.insert_val(3)
...     tree.insert_val(4)
...     tree.insert_val(5)
...     tree.insert_val(6)
...     tree.insert_val(7)
...     tree.inorder_traversal(tree.root)
...     tree.preorder_traversal(tree.root)
...     tree.postorder_traversal(tree.root)
...     tree.level_order()
...     return tree
>>> tree = TreeTest()
>>> tree.inorder == [4, 2, 5, 1, 6, 3, 7] and tree.preorder == [1, 2, 4, 5, 3, 6, 7] and tree.postorder == [4, 5, 2, 6, 7, 3, 1] and tree.levelorder == [1, 2, 3, 4, 5, 6, 7]
True
```


Tests for q3 (points: 60)
```
>>> orderProcessor = OrderProcessor()
>>> orderProcessor.timeElapsed(1)
>>> order_1 = [{'item_type':'Food','item_name':'Fries','item_size':2},
...            {'item_type':'Drink','item_name':'Coke','item_size':'small'}]
>>> orderProcessor.addOrder(order_1) 
>>> orderProcessor.timeElapsed(3)
>>> order_2 = [{'item_type':'Food','item_name':'Slider','item_size':1}]
>>> orderProcessor.addOrder(order_2) 
>>> disp1 = orderProcessor.displayQueue()
>>> orderProcessor.timeElapsed(2)
>>> order_3 = [{'item_type':'Drink','item_name':'Pepsi','item_size':'small'}] 
>>> orderProcessor.addOrder(order_3) 
>>> disp2 = orderProcessor.displayQueue()
>>> orderProcessor.timeElapsed(2)
>>> delivery1 = orderProcessor.deliverOrder()
>>> orderProcessor.timeElapsed(2)
>>> delivery2 = orderProcessor.deliverOrder()
>>> orderProcessor.timeElapsed(2)
>>> delivery3 = orderProcessor.deliverOrder()
>>> disp1 == ['order_2', 'order_1'] and disp2 == ['order_3', 'order_2', 'order_1'] and delivery1 == [['order_3', 3.99]] and delivery2 == [['order_2', 10.99]] and delivery3 == [['order_1', 23.97]]
True
```
```
>>> orderProcessor = OrderProcessor()
>>> orderProcessor.timeElapsed(1)
>>> order_1 = [{'item_type':'Food','item_name':'Fries','item_size':2}, 
...            {'item_type':'Drink','item_name':'Coke','item_size':'small'}] 
>>> orderProcessor.addOrder(order_1) 
>>> orderProcessor.timeElapsed(3)
>>> order_2 = [{'item_type':'Food','item_name':'Slider','item_size':1}] 
>>> orderProcessor.addOrder(order_2)
>>> orderProcessor.timeElapsed(2)
>>> order_3 = [{'item_type':'Drink','item_name':'Pepsi','item_size':'small'}] 
>>> orderProcessor.addOrder(order_3)
>>> orderProcessor.timeElapsed(2)
>>> d1 = orderProcessor.deliverOrder()
>>> orderProcessor.timeElapsed(4)
>>> d2 = orderProcessor.deliverOrder()
>>> d1 ==  [['order_3', 3.99]] and d2 == [['order_2', 10.99], ['order_1', 23.97]]
True
```
```
>>> order = Order('order_1',2)
>>> ans1 = order.addItem('Food','Hot Wings',2)
>>> ans2 = order.addItem('Drink','Mojito','small')
>>> ans3 = order.addItem('Drink','Coke','x-large')
>>> ans1 == 'Item Unavailable' and ans2 == 'Item Unavailable' and ans3 == 'Size Unavailable' 
True
```
