```python
# Initialize Otter
import otter
grader = otter.Notebook("assignment01.ipynb")
```

## Assignment 1

### Submission Guidelines: 
#### 1. Save and run all the cells of this notebook.
#### 2. All the assignments in this course will have two types of test cases, local test cases and hidden test cases. You will be able to test your work on local test cases by running the cell  with - grader.check_all() or grader.check("question number")
#### 3. Once you are happy with the output of grader.check_all(), you should save the notebook and upload it to gradescope
#### 4. You will be able to see the result of your work on hidden tests once the grades are released for a given assignment.

```python
import numpy as np
```

**Question 1:** Write a function `question1` that returns a list of even numbers between the range (x,y) given in arguements.<br>
Ex: x = 2, y = 7 <br>
Ans: [2,4,6]

```python
def question1(x,y):
    l=[]
    ...
    return l
    
```

```python
grader.check("q1")
```

**Question 2:** Write a function `solveMaze` that returns a solved maze.  <br>
 - Every maze (NxN) starts at (0,0) and ends at (N-1,N-1), you can traverse in empty cells (" ") but cannot travel through cells that have X in them. <br>
 - You cannot move diagonally in the maze.<br>
 - Maze has only one path and no dead ends. <br>
 - If a maze can be solved, there will be only one possible path. You need to return a solved maze that highlights this path. The path in the solved MAZE should be returned as shown below. Cells that are a part of the path should have value "-"<br>
 - If a maze  cannot be solved, return -1. <br>

Ex:<br>
[[" ","X","X"], <br>
 [" "," "," "], <br>
 ["X","X"," "]]<br>
 <br>
Ans:<br>
[["-","X","X"], <br>
 ["-","-","-"], <br>
 ["X","X","-"]]<br>

```python
maze = [["-","X","X","X"],["-","-","-","X"],["X","X","-","-"]]
print("Input: ")
print(np.array(maze))
```

```python
maze = [[" ","X","X"],[" "," "," "],["X","X"," "]]
# maze = [[" ","X","X","X"],[" "," ","X","X"],[" ","X"," "," "],[" "," ","X","X"]]
print("Input: ")
print(np.array(maze))
def solveMaze(maze):
    ...
    return maze

```

```python
grader.check("q2")
```

**Question 3:** Write a function `detectFrontRunner` that solves the following problem.  <br>
You are an analyst at the Securities and Exchange Commission agency. It is your job to find a front runner. Front-running is trading stock or any other financial asset by a broker who has inside knowledge of a future transaction that is about to affect its price substantially. A front runner is a trader that is involved in insider trading. 
You are provided with rules to detect front runners. Apply the rules to the data sample data provided and identify a front runner. Return trader_id of all the front runners (In the order the arrive) in a list.<br>

Rules:
 - A trader is flagged as a front runner if the trader has trader_type as 'ELECTRONIC'.<br>
 - There are two types of trader, electronic traders and broker traders, only electronic traders can be front runners. Broker traders cannot be front runners. <br>
 - A trade is classified as a front running trade if:
       - It occurs within 5 mins of broker trade. 
       - It is trading the same stock as the broker trade. 
       - It has the same price as the broker trade. 
       - It has the same trade_type (BUY/SELL) as the broker trade.
       - It has the same expiry date as the broker trade.

sample_data=[<br>
        {<br>
            "trader_id":"trader_12",<br>
            "trader_type":"Broker",<br>
            "trade_date":"08-31-2022"<br>
            "trade_time":"14:26:14",<br>
            "expiry":"09-23-2022",<br>
            "price":"800",<br>
            "exchange":"CBE",<br>
            "stock":"TSLA",<br>
            "trade_type":"BUY"<br>
        },<br>
    {<br>
        "trader_id":"trader_43",<br>
        "trader_type":"Electronic",<br>
        "trade_date":"08-31-2022"<br>
        "trade_time":"14:29:34",<br>
        "expiry":"09-23-2022",<br>
        "price":"800",<br>
        "exchange":"CBE",<br>
        "stock":"TSLA",<br>
        "trade_type":"BUY"<br>
    }<br>
]<br>

The output for the given sample data will be<br>
["trader_43"]


```python
data = [
    {
        "trader_id":"trader_12",
        "trader_type":"BROKER",
        "trade_date":"08-31-2022",
        "trade_time":"14:26:14",
        "expiry":"09-23-2022",
        "price":"800",
        "exchange":"CBE",
        "stock":"TSLA",
        "trade_type":"BUY"
    },
    {
        "trader_id":"trader_43",
        "trader_type":"ELECTRONIC",
        "trade_date":"08-31-2022",
        "trade_time":"14:29:34",
        "expiry":"09-23-2022",
        "price":"800",
        "exchange":"CBE",
        "stock":"TSLA",
        "trade_type":"BUY"
    },
    {
        "trader_id":"trader_23",
        "trader_type":"ELECTRONIC",
        "trade_date":"08-31-2022",
        "trade_time":"14:55:32",
        "expiry":"09-23-2022",
        "price":"800",
        "exchange":"CBE",
        "stock":"TSLA",
        "trade_type":"SELL"
    }
]
```

```python
from datetime import datetime #IGNORE
def detectFrontRunner(data):
    frontRunners=[]
    ...
    return frontRunners
detectFrontRunner(data)
```

```python
grader.check("q3")
```

**Question 4:** Write a function `generateModifiedFibonacci` that generates a modified Fibonacci sequence using recursion of length n (provided as an argument).  <br>

Fibonacci Sequence : The pattern begins after the first two numbers, 0 and 1, where each number in the sequence is always the sum of the two numbers before it.<br>
You need to identify a pattern from the samples given below and write a function to generate a sequence that follows this given pattern. <br>
Here are some examples of a ModifiedFibonacci sequence<br>
generateModifiedFibonacci(4) = [1, 2, 3, 7] <br>
generateModifiedFibonacci(6) = [1, 2, 3, 7, 22, 155] <br>
generateModifiedFibonacci(8) = [1, 2, 3, 7, 22, 155, 3411, 528706] <br>
<br>


Note: The base case for this problem is generateModifiedFibonacci(2) = [1,2]<br>
You need to solve the following problem using recursion, use of loops and other methods will result in point deduction.


```python
def generateModifiedFibonacci(n):
    ...
    return ...
def modifiedFibonacci(num):
    ...
```

```python
grader.check("q4")
```

**Question 5:** Write a function `solveMazeBackTracking` that returns a solved maze.  <br>
 - In question 1, we solve a maze that has a trivial solution, i.e. there are no dead ends in the maze provided in the first question.
 - Now that we have learned recursion, we want to apply the concept of recursion to solve mazes that have dead ends.
 - Every maze (NxN) starts at (0,0) and ends at (N-1,N-1), you can traverse in empty cells (" ") but cannot travel through cells that have X in them. <br>
 - You cannot move diagonally in the maze.<br>
 - If a maze can be solved, there will be only one possible path. You need to return 1 if a maze can be solved.<br>
 - If a maze  cannot be solved, return -1. <br>



### NOTE: There is an extra credit component in this question, if you save the solved maze that highlights the path (The path in the solved MAZE should be returned as shown below. Cells that are a part of the path should have value "-") in the solution array provided, you will receive an extra 10 points.
Ex:<br>
[[" ","X","X"], <br>
 [" "," "," "], <br>
 [" ","X"," "]]<br>
 <br>
Ans:<br>
[["-","X","X"], <br>
 ["-","-","-"], <br>
 [" ","X","-"]]<br>

```python
maze = [[" ","X","X"],[" "," "," "],[" ","X"," "]]
print("Input: ")
print(np.array(maze))

def solveMazeBackTracking(maze,display=False):
    maze = np.array(maze)
    solution = np.array(maze)
    ...
    if(display):
        return solution
    return output
print('Output:')
print(solveMazeBackTracking(maze))
```

```python
grader.check("q5")
```

### Submission Guidelines: 
#### 1. Save and run all the cells of this notebook.
#### 2. All the assignments in this course will have two types of test cases, local test cases and hidden test cases. You will be able to test your work on local test cases by running the cell  with - grader.check_all() or grader.check("question number")
#### 3. Once you are happy with the output of grader.check_all(), you should save the notebook and upload it to gradescope
#### 4. You will be able to see the result of your work on hidden tests once the grades are released for a given assignment.

```python

```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Tests for q1 (points: 10)
```
>>> question1(11,17) == [12,14,16]
True
```


Tests for q2 (points: 10)
```
>>> solveMaze([[" ","X","X"],[" ","X"," "],["X","X"," "]]) == -1
True
```
```
>>> solveMaze([[" "," ",],["X"," "]]) == [["-","-",],["X","-"]]
True
```


Tests for q3 (points: 35)
```
>>> data = [{"trader_id":"trader_12","trader_type":"BROKER","trade_date":"08-31-2022","trade_time":"14:26:14","expiry":"09-23-2022","price":"800","exchange":"CBE","stock":"TSLA","trade_type":"BUY"},
...     {"trader_id":"trader_43","trader_type":"ELECTRONIC","trade_date":"08-31-2022","trade_time":"14:29:34","expiry":"09-23-2022","price":"800","exchange":"CBE","stock":"TSLA","trade_type":"BUY"},
...     {"trader_id":"trader_23","trader_type":"ELECTRONIC","trade_date":"08-31-2022","trade_time":"14:55:32","expiry":"09-23-2022","price":"800","exchange":"CBE","stock":"TSLA","trade_type":"SELL"}]
>>> detectFrontRunner(data) == ['trader_43']
True
```


Tests for q4 (points: 15)
```
>>> generateModifiedFibonacci(4) == [1, 2, 3, 7]
True
```


Tests for q5 (points: 35)
```
>>> solveMazeBackTracking([[" ","X","X"],[" ","X"," "],["X","X"," "]]) == -1
True
```
```
>>> solveMazeBackTracking([[" "," ",],["X"," "]]) == 1
True
```
