```python
# Initialize Otter
import otter
grader = otter.Notebook("assignment00.ipynb")
```

## Assignment 0

### This is assignment 0. The purpose of this assignment is to help you get familiar with assignment notebooks, datahub and gradescope submission.

**Question 1:** Write a function `add` that returns the addition of its argument.
Test your solution by running the cells below function definition.

```python
def add(a,b):
    ...
```

```python
grader.check("q1")
```

### Submission Guidelines: 
#### 1. Save and run all the cells of this notebook.
#### 2. All the assignments in this course will have two types of test cases, local test cases and hidden test cases. You will be able to test your work on local test cases by running the cell  with - grader.check_all() or grader.check("question number")
#### 3. Once you are happy with the output of grader.check_all(), you should save the notebook and upload it to gradescope
#### 4. You will be able to see the result of your work on hidden tests once the grades are released for a given assignment.

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Tests for q1 (points: 10)
```
>>> add(1,2) == 3
True
```
