```python

```

```python
# Initialize Otter
import otter
grader = otter.Notebook("midterm.ipynb")
```

## Midterm
### You have 60 minutes to solve 4 questions given below. 

### Do not change the code that has been provided to you. You can add code but do not change/remove any components provided by us.

```python
import numpy as np
import pandas as pd
```

**Question 1:** Complete the function `question1` based on the instructions given below. <br>

### Given a numpy array, return the maximum element from each row in a list

x=
[[4 8 3]<br>
  [2 9 7]<br>
  [0 2 7]]<br>

question1(x) = [8,9,7]

```python
def question1(arr):
    '''
    1 local test case : 5 points
    1 hidden test case : 5 points
    '''
    l=[]
    for x in arr:
        l.append(max(x))
    return l
```

```python
grader.check("q1")
```

Output:
```
q1 results: All test cases passed!
```

**Question 2:** Complete the function `question2` based on the instructions given below. <br>

### Given an integer argument return factorial of the value. Hint: Use recursion.

Note: The factorial function says to multiply all whole numbers from our chosen number down to 1.

Examples:

question4(4) = 4 × 3 × 2 × 1 = 24


```python
def question2(num):
    '''
    1 local test case : 5 points
    1 hidden test case : 5 points
    '''
    ans = 1
    ans = fact(num)
    return ans

def fact(x):
    if (x<=1):
        return 1
    return x*fact(x-1)
```

```python
grader.check("q2")
```

Output:
```
q2 results: All test cases passed!
```

**Question 3:** Complete the function `question3` based on the instructions given below.<br>

### Apply the following instructions (in order) to an input two dimensional numpy array with 0s and 1s and return the output


**\<Operations to perform\>**

<ol>
    <li> Reverse the odd rows.
    <li> Once the odd rows are reversed, replace the first element of every column according to the following logic : If the column contains even number of 1's replace the first element of that column with a 0. If a column has odd number of 1's replace the first element of that column with a 1.<br>
        Ex: column = [1, 0, 0, 1]<br>
        numberOf1s(column) = 2; even, so we replace the first element of the column with a 0<br>
        column = [0, 0, 0, 1]
    <li> Replace the 0's in each column by the mean of the array.<br>
        Ex. col = [1, 0, 0, 1]<br>
        array_mean = 0.5; here the array_mean is the mean of the entire array.<br>
        output = [1, 0.5, 0.5, 1]
    <li> Multiply all the elements of the array with 255
    <li> Return the transpose of the array
    
</ol>

```python
arr2 = np.array([[1,0,0], [1,1,0], [0,0,0], [1,0,1]])
np.where(arr2==0, 122, arr2)
arr2
```

Output:
```
array([[1, 0, 0],
       [1, 1, 0],
       [0, 0, 0],
       [1, 0, 1]])
```

```python
def question3(arr):
    '''
    1 local test case : 15 points
    1 hidden test case : 25 points
    '''
    #1
    for x in range(len(arr)):
        if x%2==1:
            arr[x] = arr[x][::-1]
    #2 Go over columns
    for x in range(len(arr[0])):
        if numberOf1sInColumn(arr, x)%2==0:
            arr[0][x] = 0
        else:
            arr[0][x] = 1
    #3
    meanValue = arr.mean()
    arr = np.where(arr==0, meanValue, arr)
    #4
    arr = arr*255
    #5
    arr = arr.T      
    return arr

def numberOf1sInColumn(arr, columnNo):
    count = 0
    for i in range(len(arr)):
        if arr[i][columnNo]==1:
            count+=1
    return count

question3(np.array([[1,0,0], [1,1,0], [0,0,0], [1,0,1]]))
```

Output:
```
array([[106.25, 106.25, 106.25, 255.  ],
       [255.  , 255.  , 106.25, 106.25],
       [106.25, 255.  , 106.25, 255.  ]])
```

```python
grader.check("q3")
```

Output:
```
q3 results: All test cases passed!
```

**Question 4:** Complete the function `question4` based on the instructions given below.<br>

### Convert a given number to its Roman counter part in a pandas DataFrame.

Input numbers range from 1 to 50 only, You are given these numbers in a column in a pandas dataframe. You need to create a new column in the SAME dataframe with column name <b>"Roman"</b> and list the corresponding Roman counterparts in that column. 

Roman numerals are represented by seven different symbols: I, V, X, and  L<br>
Symbol:Value<br>
I      :       1<br>
V       :      5<br>
X        :     10<br>
L         :    50<br>
For example, 2 is written as II in Roman numeral, just two one's added together. 12 is written as XII, which is simply X + II. The number 27 is written as XXVII, which is XX + V + II.<br>
Roman numerals are usually written largest to smallest from left to right. However, the numeral for four is not IIII. Instead, the number four is written as IV. Because the one is before the five we subtract it making four. The same principle applies to the number nine, which is written as IX. There are six instances where subtraction is used:<br>
I can be placed before V (5) and X (10) to make 4 and 9.<br>
X can be placed before L (50) to make 40.<br>
Given an integer in the range of 1 to 50 (both inclusive), convert it to a roman numeral.<br>
### NOTE: If you need a Roman mapping to any numeric value, you can ask us and we will provide it to you. You need to use df.apply to create the new Roman column. We have provided you with a function that you can use with df.apply


```python
'''
This is an example of how the input will look like
'''
df_dict = {'Numeric':[4,12,11]}
inputDataFrame = pd.DataFrame(df_dict)
inputDataFrame

```

Output:
```
Numeric
0        4
1       12
2       11
```

```python
'''
This is an example of how the expected output dataframe should look like.
'''
df_dict = {'Numeric':[4,12,11],'Roman':['IV','XII','XI']}
outputDataFrame = pd.DataFrame(df_dict)
outputDataFrame
```

Output:
```
Numeric Roman
0        4    IV
1       12   XII
2       11    XI
```

```python
def convertNumeric2Roman(numeric_value):
    '''
    We have provided this method to be used with df.apply
    '''
    romanValue=""
    numberMapping = {1:"I", 5:"V", 10:"X", 50:"L"}
    while(numeric_value>0):
        digit = numeric_value%10;
        
        if digit in numberMapping.keys():
            romanValue += numberMapping[digit]
        elif (digit<4):
            romanValue += numberMapping[digit]*digit
        elif (digit==4):
            romanValue += numberMapping[1] + numberMapping[5]
        elif (digit<9):
            romanValue += numberMapping[5] + numberMapping[1]*(digit-5)
        elif (digit==9):
            romanValue += numberMapping[5] + numberMapping[10]
        elif (digit<14):
            romanValue += numberMapping[5] + numberMapping[1]*(digit-5)
        else:
            pass
        numeric_value = numeric_value/10
    return romanValue

def question4(dataFrame): 
    '''
    1 local test case : 15 points
    1 hidden test case : 25 points
    '''
    dataFrame["Roman"] = dataFrame["Numeric"].apply(convertNumeric2Roman)
    return dataFrame
question4(inputDataFrame)
```

Output:
```
File ~\anaconda3\lib\site-packages\pandas\_libs\lib.pyx:2859, in pandas._libs.lib.map_infer()

Input In [16], in convertNumeric2Roman(numeric_value)
     11     romanValue += numberMapping[digit]
     12 elif (digit<4):
---> 13     romanValue += numberMapping[digit]*digit
     14 elif (digit==4):
     15     romanValue += numberMapping[1] + numberMapping[5]

KeyError: 0.4
```

```python
grader.check("q4")
```

### Submission Guidelines: 
#### 1. Save and run all the cells of this notebook.
#### 2. All the assignments in this course will have two types of test cases, local test cases and hidden test cases. You will be able to test your work on local test cases by running the cell  with - grader.check_all() or grader.check("question number")
#### 3. Once you are happy with the output of grader.check_all() and upload the notebook to gradescope
#### 4. You will be able to see the result of your work on hidden tests once the grades are released for a given assignment.

```python

```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Tests for q1 (points: 10)
```
>>> question1(np.array([[3, 7, 8],[6, 8, 8],[9, 2, 6]])) == [8,8,9]
True
```


Tests for q2 (points: 10)
```
>>> question2(5) == 120
True
```


Tests for q3 (points: 40)
```
>>> x = np.array([[1, 0, 0],[0, 1, 1],[1, 0, 0],[1, 1, 0]])
>>> flags = question3(x) == np.array([[255.  , 255.  , 255.  , 148.75],[148.75, 255.  , 148.75, 255.  ],[255.  , 148.75, 148.75, 255.  ]])
>>> sum(sum(flags)) == 12
True
```


Tests for q4 (points: 40)
```
>>> df_dict = {'Numeric':[4,9,14]}
>>> inputDataFrame = pd.DataFrame(df_dict)
>>> list(question4(inputDataFrame)['Roman']) == ['IV', 'IX', 'XIV']
True
```
