```python
# Initialize Otter
import otter
grader = otter.Notebook("assignment02.ipynb")
```

## Assignment 2

### Do not change the code that has been provided to you. You can add code but do not change/remove any components provided by us.

```python
import numpy as np
import pandas as pd
import matplotlib.pyplot as plt
```

```python
data = pd.read_csv('./data.csv')
data.head()
```

**Question 1:** Write a function `data_preprocessing` that returns a preprocessed data based on the rules given below.<br>

### You need to clean the data to prepare it for the next steps: Follow the following steps.
<ul>
    <li>Shipment column can only have three types : {'Flight', 'Road', 'Ship'}. Any row that does not have the given shipment type, should be removed.</li>
    <li>The Rating column can have the following values : {1.0,2.0,3.0,4.0,5.0}. Any row that has a rating value other than the given values, should be replaced by the average value of the rating column</li>
    <li>The Weight column has numeric values. Any row that has a weight value other than a valid numeric value, should be replaced by the average value of the weight column</li>
    <li>The Discount column can have numeric values in the range (0 to 80). Any row that has a weight value other than a valid value, should be replaced by the average value of the discount column</li>
    <li>Any row that has the Shipment Delivered (1/0) value other than 1 or 0 should be removed.</li>
</ul>



```python
def data_preprocessing(data):
    ...
    return data
data = data_preprocessing(data)
```

```python
grader.check("q1")
```

<!-- BEGIN QUESTION -->

**Question 2:** Write a function `data_analysis_matplotlib` that displays plots based on the requirements given below.<br>

### Display plots based on the requirements given below.
<ul>
    <li>Display a plot that has two subplots, the first subplot should be a grouped bar chart between number of shipments delivered and Product_importance: groups(shipment delivered/ shipment not delivered), the second subplot should be a grouped bar chart between the number of shipments delivered and Region: groups(shipment delivered/ shipment not delivered) </li>
</ul>
The plot should look like this:<br>
<img src='plot.png' width = "800" height="400">
Note : You might have to find out the different categories and range of columns in the data to come up with a more effective visualization.


```python
def data_analysis_matplotlib(data):
    plt.figure(figsize=(16,8))
    
    ...
    plt.show()
data_analysis_matplotlib(data)
```

<!-- END QUESTION -->

**Question 3:** Write a function `one_hot_encoding` that encodes feature vector according to details given below.<br>

### One-hot Encoding data for machine learning.
<ul>
    <li>We need to create a feature vector (X) and label vector (y) that we will use in our Machine Learning Model.</li>
    <li>In order to create a feature vector, we need to one-hot encode the categorical features - Region, Shipment, Product_importance, Gender</li>
</ul>


```python
data.head()
```

```python
def one_hot_encoding(data):
    data['Gender'] = data['Gender'].replace(to_replace = "F", value = "Female")
    data['Gender'] = data['Gender'].replace(to_replace = "M", value = "Male")
    
    categorical_feats = ['Region','Shipment','Product_importance','Gender']
    ...
    return data
one_hot_data = one_hot_encoding(data)
X = np.array(one_hot_data.drop('Shipment Delivered (1/0)', axis=1))
y = np.array(one_hot_data['Shipment Delivered (1/0)'])
```

```python
grader.check("q3")
```

```python
import json
data={}
data['X_train']=X[:9000].tolist()
data['y_train']=y[:9000].tolist()
data['X_test']=X[9000:].tolist()
data['y_test']=y[9000:].tolist()
```

```python
with open("data.json", "w") as outfile:
    json.dump(data, outfile)
```

```python
'''
If you get the error: ModuleNotFoundError: No module named 'sklearn',
uncomment the line below and run this cell, then restart the notebook and run all cells again
'''
# !pip install scikit-learn
```

**Question 4:**<br>

### Run the machine learning model in model.py
<ul>
    <li>You have been provided a model.py file along with this notebook.</li>
    <li>There are two calls stated below, one for training and other for testing. Both calls differ in the system arguements that are passed to the main function. </li>
<li>If -train is passed as a system argument, then you need to run_model function as run_model("data.json",train = True) 
    </li>
<li>If -test is passed as a system argument, then you need to run_model function as run_model("data.json",train = False) 
    </li>
</ul>


```python
!python model.py -train
```

```python
!python model.py -test
```

```python
import json
with open("output_train.json", "r") as outfile:
    output_train=json.load(outfile)
with open("output_test.json", "r") as outfile:
    output_test=json.load(outfile)
```

```python
grader.check("q4")
```

**Question 5:**<br>

### Regex
<ul>
    <li>Complete the functions given below to provide regex functionality</li>
</ul>


```python
import re
def regex1(string):
    '''
    Given an IP number as a string remove all leading zeros in the IP address
    Ex: regex1("123.01.09.103") = '123.1.9.103'
    
    '''
    ans = None
    ...
    return ans
regex1("123.01.09.103")
```

```python
def regex2(string):
    '''
    Given an email return True if the email is valid, else return False
    regex2('john.doe-gmail.com') = False
    '''
    flag = None
    ...
    return flag
regex2('john.doe-gmail.com')
```

```python
grader.check("q5")
```

### Submission Guidelines: 
#### 1. Save and run all the cells of this notebook.
#### 2. All the assignments in this course will have two types of test cases, local test cases and hidden test cases. You will be able to test your work on local test cases by running the cell  with - grader.check_all() or grader.check("question number")
#### 3. Once you are happy with the output of grader.check_all(), you should save the following files in a zip to gradescope:  assignment02.ipynb, model.py, output_train.json, output_test.json, data.csv
#### 4. You will be able to see the result of your work on hidden tests once the grades are released for a given assignment.

```python

```

```python

```

```python

```

```python

```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Tests for q1 (points: 30)
```
>>> data.isnull().sum().sum()==0
True
```


Tests for q3 (points: 20)
```
>>> X.shape == (10618, 20)
True
```
```
>>> y.shape == (10618,)
True
```
```
>>> arr = np.array([1.000e+00, 4.000e+00, 2.000e+00, 1.770e+02, 3.000e+00, 4.400e+01, 1.233e+03,0.000e+00, 0.000e+00, 0.000e+00, 1.000e+00, 0.000e+00, 1.000e+00, 0.000e+00, 0.000e+00, 0.000e+00, 1.000e+00, 0.000e+00, 1.000e+00, 0.000e+00])
>>> base = X[0],
>>> sum(sum(np.equal(arr,base))) == 20
True
```


Tests for q4 (points: 10)
```
>>> output_train['training_accuracy'] == 1.0
True
```
```
>>> output_test['training_accuracy'] == 1.0 and round(output_test['testing_accuracy'],2) == 0.54
True
```


Tests for q5 (points: 10)
```
>>> regex1("103.012.32.87") == "103.12.32.87" and regex1("100.001.100.001") == "100.1.100.1"
True
```
```
>>> regex2('john.doe-gmail.com') == False and regex2('jack.fro@gmail.com') == True
True
```
