```python
# Initialize Otter
import otter
grader = otter.Notebook("assignment04.ipynb")
```

## Assignment 4

### Do not change the code that has been provided to you. You can add code but do not change/remove any components provided by us.

```python
import numpy as np
```

**Question 1:** Complete the Class `Graph` that provides functionality based on rules given below.<br>

### Graph
<ul>
    <li>Initialize a graph with an empty dictionary.</li>
    <li>Implement class functions based on the comments given in them.</li>
</ul>


```python
class Graph:
    def __init__(self):
        '''
        Initialize graph with empty dictionary
        A graph will be contructed using a dictionary where a key is a node and its value is a list of nodes which are the child nodes of the key node
        graph = {
            key_node : [child_node_1, child_node_2, child_node_3]
        }
        '''
        self.graph = None
        ...
    def display_graph(self):
        print(self.graph)
    def add_edge(self,parent_node, child_node):
        '''
        Adds a child node to the list of child nodes of the parent node
        Ex: graph = {
            3 : [2,4]
        }
        add_edge(3,5) -> graph = {
            3 : [2,4,5]
            5 : [3]
        }
        '''
        ...
    def dfs(self,source_node, destination_node):
        '''
        Traverse the graph using DFS and return all the paths from source_node to destination_node
        '''
        paths = []
        ...
        return paths
    def bfs(self,source_node, destination_node):
        '''
        Traverse the graph using BFS and return all the paths from source_node to destination_node
        '''
        paths = []
        ...
        return paths

```

```python
grader.check("q1")
```

**Question 2:** Foreign Exchange rate provider.<br>

### Based on the the instructions given below design a Forex System.

<ul>
    <li>You are provided with a class Forex. You need to fill in the methods to make the Forex system function as required. </li>
    <li>Each method has a set of comments which will tell you what is expected as an output of that method.
    <li>You will need to maintain a graph in the Forex task to calculate the exchange rate. You are provided with reference dictionaries to calculate estimated exchange rate.</li>
    <li>Given the complexity of this task, we will not have any hidden test case for this question.</li>
</ul>


Hint: We highly recommend using the Graph data structure and its methods we contructed in the previous question

### Notes on constructing the graph.  You need to contruct a bi-directional graph. Here are some tips on constructing the bidirectional graph.

<ul>
    <li>You need a map/dictionary to story edge values, for example given a key value pair 'USDCAD':1.33, USD and CAD will be the nodes in this case using which we will construct a graph.</li>
    <li>Now coming to storing edge values, given the above key value pair, we can store two edge values, USD -> CAD = 1.33 and CAD -> USD = 1/1.33.</li>
</ul>


```python
class Forex:
    def __init__(self, reference_rates):
        '''
        Initialize a graph to handle incoming reference rates.
        The format of reference_rates is given below:
        reference_rates = {
            'USDCAD':1.33,
            'USDEUR':0.96,            
        }
        '''
        self.graph = None
        self.reference_rates = None
        ...
    def construct_graph(self):
        '''
        Construct a graph based on the reference rates provided in the constructor
        
        Ex: graph = {
            'USD': ['CAD','EUR'],
            'CAD': ['USD'],
            ''
        }
        '''
        
        ...
    def get_exchange_rate(self,source,destination):
        '''
        Return the exchange rate from source to destination with 2 floating points
        
        Ex: self.get_exchange_rate('CAD','EUR') = 0.72
        
        '''
        ...
```

```python
'''
Here is a small example so you understand the implementation
How to handles edges 1.33 and 1/1.33
'''
reference_rates = {'USDCAD':1.33,'USDEUR':0.96}
forex = Forex(reference_rates)
forex.construct_graph()
exchange_rate = forex.get_exchange_rate('CAD','EUR')
# print(exchange_rate) #prints 0.72
```

```python
grader.check("q2")
```

### Submission Guidelines: 
#### 1. Save and run all the cells of this notebook.
#### 2. All the assignments in this course will have two types of test cases, local test cases and hidden test cases. You will be able to test your work on local test cases by running the cell  with - grader.check_all() or grader.check("question number")
#### 3. Once you are happy with the output of grader.check_all(), you should save the following files in a zip to gradescope:  assignment04.ipynb  (Make sure the name of the files is consistent when you upload it to gradescope)
#### 4. You will be able to see the result of your work on hidden tests once the grades are released for a given assignment.

```python

```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Tests for q1 (points: 20)
```
>>> graph = Graph()
>>> graph.add_edge(1,2)
>>> graph.add_edge(1,3)
>>> graph.add_edge(2,3)
>>> graph.add_edge(3,5)
>>> graph.add_edge(2,5)
>>> paths_dfs = graph.dfs(1,5)
>>> # paths_dfs
>>> [1,2,5] in paths_dfs and [1,3,5] in paths_dfs and [1,2,3,5] in paths_dfs
True
```
```
>>> graph = Graph()
>>> graph.add_edge('A','B')
>>> graph.add_edge('A','D')
>>> graph.add_edge('B','D')
>>> graph.add_edge('D','E')
>>> graph.add_edge('C','E')
>>> graph.add_edge('B','C')
>>> graph.add_edge('C','D')
>>> 
>>> paths = graph.bfs('A','E')
>>> ['A', 'D', 'E'] in paths and ['A', 'B', 'D', 'E'] in paths and  ['A', 'B', 'C', 'E'] in paths and ['A', 'B', 'C', 'D', 'E'] in paths
True
```


Tests for q2 (points: 30)
```
>>> reference_rates = {'USDCAD':1.33,'USDEUR':0.96,'EURJPY':145.08, 'CADINR':61.26, 'AUDGBP':0.56}
>>> forex1 = Forex(reference_rates)
>>> forex1.construct_graph()
>>> forex1.get_exchange_rate('INR','JPY') == 1.71 and forex1.get_exchange_rate('AUD','JPY') == -1
True
```
```
>>> reference_rates = {'USDCAD':1.33,'USDEUR':0.96,'EURJPY':145.08, 'CADINR':61.26, 'EURGBP':0.87, 'CADJPY':104.98, 'GBPJPY':166.40}
>>> forex2 = Forex(reference_rates)
>>> forex2.construct_graph()
>>> ans2_1 = forex2.get_exchange_rate('GBP','INR')
>>> ans2_2 = forex2.get_exchange_rate('USD','JPY')
>>> round(ans2_1,1) == 97.1 and round(ans2_2,1) ==  139.6
True
```
